Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **37 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Naveen"
COLLABORATORS = "none"
```

---

# CS110 Pre-class Work - Asymptotic notations; maximum-subarray problem

## Question 1 (Exercise 3.1-3 of Cormen, et al. ) [time estimate: 5 minutes]
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is true but meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
This statement says that for worst case scenario, the running time for algorithm A will be O(n^2). O(n^2) provides the upper bound of a function whereas the statement says at least thus giving a lower bound. This two terms can be contradictory and make the statement meaningless since A can have any possible run time. 

## Question 2 (Exercise 3.1-4 of Cormen, et al. ) [time estimate: 5 minutes]

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$? Why?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
Assuming, f(n) is O(g(n)) if a constant multiple of g is bigger than or equal to f. 
2^(n+1) is O(2n.
A constant multiple of 2n is always greater than 2n+1.
Multiply 2n by 2 = 2n+1, which is is greater than or equal to 2n+1!
Hence, with c=2 and k=1, we have 2×2n ≥ 2n+1 for all  n ≥ 1. Thus, 2n+1 is O(2n).

## Question 3 [time estimate: 20 minutes]
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
x=input("enter a list of numbers")
list1=x.split()
n=len(list1)
for f in range(n):
  list1[f]=int(list1[f])
  mx=0
for i in range(n):
  for j in range (i+1,n):
    s=0
    for k in range(i,j+1):
      s=list1[k]+s
      if s>list1[k]:
        mx=s
```

Output:
```
enter a list of numbers3,47,6,8
ValueError                                Traceback (most recent call last)
<ipython-input-14-d2802c5a0f13> in <module>
      3 n=len(list1)
      4 for f in range(n):
----> 5   list1[f]=int(list1[f])
      6   mx=0
      7 for i in range(n):

ValueError: invalid literal for int() with base 10: '3,47,6,8'
```

```python
def find_max_subarray_brute(array) #input array list, used brute force to find maximum subarray algorithm 

    int max, start, end; #maximum integer value element from start till end, range defined
    for (int i = 0; i < array.size(); i++) 
            
        int sum = 0;
        for (int j = i; j < array.size(); j++) 
        
            sum += array[j];        
   
    if (sum > max) #if sum of subarray is maximum then terminate loop and return ith and jth values/elements in the subarray
                max = sum;
                start = i; #starting element
                end = j; #ending element
                            
return [max: max, start: start, end: end]



```

Output:
```
File "<ipython-input-25-c8aa98f2aa5d>", line 1
    def find_max_subarray_brute(array) #input array list, used brute force to find maximum subarray algorithm
                                                                                                              ^
SyntaxError: invalid syntax
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-12-5ddab6f2f574> in <module>
----> 1 assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
      2 assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))

NameError: name 'bruteforce_max_subarray' is not defined
```

## Question 4 [time estimate: 2 minutes]
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 5. Asymptotic notation [time estimate: 5 minutes]
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Running time 
--- | --- 
Insertion sort | 
Selection sort |  
Bubble sort |  
Finding maximum subarray |  

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
Algorithm | Running time 
--- | --- 
Insertion sort | Θ(n)
Selection sort |  O(n^2)
Bubble sort |  
Finding maximum subarray |  

## [Optional] Question 6 [time estimate: 15 minutes] 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """
    Implements brute-force minimum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the minimum subarray
        - the end index of the minimum subarray
        - the value of the minimum subarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```