Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **50 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Naveen"
COLLABORATORS = "none"
```

---

# CS110 Pre-class Work - Analyzing elementary sorting algorithms



## Question 1 [time estimate: 10 minutes]
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
Line 5 till 8 show the while loop and i represents the ith element in the array A. In line 7, 1 is subtracted from the element in order to preserve the loop invariant as we added 1 in line 6 to i. 

## Question 2 [time estimate: 15 minutes]
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertion_sort(A):
    for j in range(2, len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertion_sort(A):
    for j in range(2, len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertion_sort([0]) == [0])
assert(insertion_sort([-1,1]) == [-1,1])
assert(insertion_sort([1,-1]) == [-1,1])
assert(insertion_sort([1,6,3,6]) == [1,3,6,6])
assert(insertion_sort([7, 3, 5, 4, 6, 4]) == [3, 4, 4, 5, 6, 7])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-19-951b635b7173> in <module>
      1 assert(insertion_sort([0]) == [0])
      2 assert(insertion_sort([-1,1]) == [-1,1])
----> 3 assert(insertion_sort([1,-1]) == [-1,1])
      4 assert(insertion_sort([1,6,3,6]) == [1,3,6,6])
      5 assert(insertion_sort([7, 3, 5, 4, 6, 4]) == [3, 4, 4, 5, 6, 7])

AssertionError:
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
Assertion error was used for line 3 as it showed insertion sort 1,-1= -1,1. i[-1] represented the last element in the list so it was made to show A=1 which shows first element in python list. or it can also be changed to A=1,-1 == 1,-1.

## Question 3 [time estimate: 15 minutes]
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertion_sort`,  your new Python function, `bubble_sort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubbleSort( seqa ): #defines functions and accepts parameter seqa, help taken through online resources for the code
    n = len( seqa  ) #length of array seqa 

    for i in range( n - 1 ) : #a for loop is run with (n-1) times 
        flag = 0 #flag used for optimisation of code to ensure if bubble sorting algorithm has made a swap or not

        for j in range(n - 1) : #inner loop started for the algorithm, comparing values
            
            if seqa[j] > seqa[j + 1] : #if statement used to check values of LHS and RHS 
                tmp = seqa[j]
                seqa[j] = seqa[j + 1]
                seqa[j + 1] = tmp
                flag = 1 #one swap has happened 

        if flag == 0:  #variable of flag checked
            break #termination condition for inner loop 

    return seqa #return list after sorting it 

ab = [22,8,6,11,2] #define variable ab 

result = bubbleSort(ab) #list run through the algorithm 

print (result) #result is printed
```

Output:
```
[2, 6, 8, 11, 22]
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
The three test cases could be another sample test with a random list of elements, another could be an example of extensive testing with load testing by increasing the input value to check the efficiency of the algorithm or find a case where the code has no solution.

## Question 4 [time estimate: 10 minutes]
The following pseudocode corresponds to a third sorting algorithm covered in this lesson, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free by providing at least three test cases. `selection_sort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selection_sort(A): 

   for i in range(len(A)):

      # it finds the minimum element in remaining list since, help taken from online resources
       minPosition = i

       for j in range(i+1, len(A)):
           if A[minPosition] > A[j]:
               minPosition = j
                
       # it swaps the found minimum element with minimum element position in the list      
       temp = A[i]
       A[i] = A[minPosition]
       A[minPosition] = temp

   return A

print(selection_sort([3,4,33,91,0,45,62]))
print(selection_sort([21,3,9,10,0,4,11]))
print(selection_sort([2,0,33,91,1,98,89]))
```

Output:
```
[0, 3, 4, 33, 45, 62, 91]
[0, 3, 4, 9, 10, 11, 21]
[0, 1, 2, 33, 89, 91, 98]
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**